Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Conditioning

Conditioning of a function $f(x)$ asks about the behavior of $f$ for points $x + \delta x$ "nearby" to $x$.  When "nearby" is measured in absolute terms, as in
$$\lvert \delta x \rvert < \text{tolerance} $$
then we will care about the absolute condition number
$$\hat \kappa = \lvert f'(x) \rvert .$$
For floating point arithmetic, we evaluate closeness as
$$ \left\lvert \frac{\delta x}{x} \right\rvert < \text{tolerance} $$
and thus are interested in the relative condition number
$$ \kappa = |f'(x)| \frac{|x|}{|f|} .$$

```python
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('ggplot')
import numpy as np

def plus_one(x):
    return x + 1

def plot_abscond(f, x, tol):
    xx = x + tol*np.random.randn(100);
    yy = f(xx)
    plt.plot(xx, yy, 'o')
    plt.xlabel('x')
    plt.ylabel('y')
    plt.title('Absolute conditioning')
    return xx, yy;
    
plot_abscond(plus_one, 1, 1e-2);
```

The absolute condition number is seen in the slope of this line.

```python
def plot_relcond(f, x, tol):
    # 1e-16 to avoid dividing by zero if x=0
    xscale = 1e-16 + np.abs(x)
    yscale = 1e-16 + np.abs(f(x))
    xx = x + tol * xscale * np.random.randn(100);
    yy = f(xx)
    plt.plot(xx/xscale, yy/yscale, 'o')
    plt.xlabel('x/xscale')
    plt.ylabel('y/yscale')
    plt.title('Relative conditioning')
    return xx/xscale, yy/yscale
    
plot_relcond(plus_one, 1, 1e-2);
```

The relative condition number is seen in the slope.

```python
def f(x):
    return np.log(1e-10 + np.abs(x - 0.5))
```

<!-- grade_id: cell-8dc2d645f9359607, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-010bcebcc4bd7e9b, points: 1, grade: true, solution: false -->
```python
xx, yy = plot_abscond(f, bigabs(), 1e-4)
assert yy.max() - yy.min() > 1, "Try a different point"
```

```python
def g(x):
    return np.log(1e-10 + np.abs(x - 0.5))
```

<!-- grade_id: cell-14f6036919b7a25e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-525a0e77a5444dfe, points: 1, grade: true, solution: false -->
```python
xx, yy = plot_relcond(g, bigrel(), 1e-4)
assert yy.max() - yy.min() > 1000, "Try a different x"
```